---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Introduction to SpeechRecognition Python Library

```yaml
type: VideoExercise
key: c2215db811
xp: 50
```

`@projector_key`
bc1a09d45f5a3931105ea89e3852664e
